Jun  6 05:08:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check users-login;WARNING;SOFT;1;USERS WARNING - 1 users currently logged in
Jun  6 05:09:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check users-login;WARNING;SOFT;2;USERS WARNING - 1 users currently logged in
Jun  6 05:10:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check users-login;WARNING;HARD;3;USERS WARNING - 1 users currently logged in
Jun  6 05:52:55 mgmt-b nagios3: SERVICE ALERT: backup-b.foo.org.nz;check users-login;WARNING;SOFT;1;USERS WARNING - 1 users currently logged in
Jun  6 05:53:55 mgmt-b nagios3: SERVICE ALERT: backup-b.foo.org.nz;check users-login;WARNING;SOFT;2;USERS WARNING - 2 users currently logged in
Jun  6 05:54:55 mgmt-b nagios3: SERVICE ALERT: backup-b.foo.org.nz;check users-login;WARNING;HARD;3;USERS WARNING - 2 users currently logged in
Jun  6 06:21:05 mgmt-b nagios3: SERVICE ALERT: db-b.foo.org.nz;check users-login;WARNING;SOFT;1;USERS WARNING - 2 users currently logged in
Jun  6 06:22:05 mgmt-b nagios3: SERVICE ALERT: db-b.foo.org.nz;check users-login;WARNING;SOFT;2;USERS WARNING - 2 users currently logged in
Jun  6 06:23:05 mgmt-b nagios3: SERVICE ALERT: db-b.foo.org.nz;check users-login;WARNING;HARD;3;USERS WARNING - 2 users currently logged in
Jun  6 06:44:05 mgmt-b nagios3: SERVICE ALERT: db-b.foo.org.nz;MariaDB servers;CRITICAL;SOFT;1;Access denied for user 'nagios'@'mgmt-b.foo.org.nz' (using password: YES)
Jun  6 06:45:05 mgmt-b nagios3: SERVICE ALERT: db-b.foo.org.nz;MariaDB servers;CRITICAL;SOFT;2;Access denied for user 'nagios'@'mgmt-b.foo.org.nz' (using password: YES)
Jun  6 06:46:05 mgmt-b nagios3: SERVICE ALERT: db-b.foo.org.nz;MariaDB servers;CRITICAL;HARD;3;Access denied for user 'nagios'@'mgmt-b.foo.org.nz' (using password: YES)
Jun  6 07:05:59 mgmt-b nagios3: SERVICE ALERT: db-b.foo.org.nz;MariaDB servers;OK;HARD;3;Uptime: 270176  Threads: 3  Questions: 40545  Slow queries: 0  Opens: 169  Flush tables: 1  Open tables: 160  Queries per second avg: 0.150
Jun  6 07:23:54 mgmt-b nagios3: SERVICE ALERT: localhost;check owncloud database test;CRITICAL;SOFT;1;CRITICAL - cannot connect to owncloudtest. Unknown database 'owncloudtest'
Jun  6 07:24:54 mgmt-b nagios3: SERVICE ALERT: localhost;check owncloud database test;CRITICAL;SOFT;2;CRITICAL - cannot connect to owncloudtest. Unknown database 'owncloudtest'
Jun  6 07:25:57 mgmt-b nagios3: SERVICE ALERT: localhost;check owncloud database test;UNKNOWN;HARD;3;NRPE: Command 'check_owncloud_db_test' not defined
Jun  6 07:51:05 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check puppet-nodes;UNKNOWN;SOFT;1;NRPE: Command 'check_puppet_nodes' not defined
Jun  6 07:51:55 mgmt-b nagios3: SERVICE ALERT: backup-b.foo.org.nz;check puppet-nodes;UNKNOWN;SOFT;1;NRPE: Command 'check_puppet_nodes' not defined
Jun  6 07:52:05 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check puppet-nodes;UNKNOWN;SOFT;2;NRPE: Command 'check_puppet_nodes' not defined
Jun  6 07:52:55 mgmt-b nagios3: SERVICE ALERT: backup-b.foo.org.nz;check puppet-nodes;UNKNOWN;SOFT;2;NRPE: Command 'check_puppet_nodes' not defined
Jun  6 07:53:05 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check puppet-nodes;OK;SOFT;3;Active: active (running) since Mon 2022-06-06 07:52:47 UTC: 9s ago
Jun  6 07:53:35 mgmt-b nagios3: SERVICE ALERT: db-b.foo.org.nz;check puppet-nodes;UNKNOWN;SOFT;1;NRPE: Command 'check_puppet_nodes' not defined
Jun  6 07:53:55 mgmt-b nagios3: SERVICE ALERT: backup-b.foo.org.nz;check puppet-nodes;UNKNOWN;HARD;3;NRPE: Command 'check_puppet_nodes' not defined
Jun  6 07:54:35 mgmt-b nagios3: SERVICE ALERT: db-b.foo.org.nz;check puppet-nodes;UNKNOWN;SOFT;2;NRPE: Command 'check_puppet_nodes' not defined
Jun  6 07:55:35 mgmt-b nagios3: SERVICE ALERT: db-b.foo.org.nz;check puppet-nodes;UNKNOWN;HARD;3;NRPE: Command 'check_puppet_nodes' not defined
Jun  6 07:58:55 mgmt-b nagios3: SERVICE ALERT: backup-b.foo.org.nz;check puppet-nodes;OK;HARD;3;Active: active (running) since Mon 2022-06-06 07:57:36 UTC: 1min 13s ago
Jun  6 08:00:35 mgmt-b nagios3: SERVICE ALERT: db-b.foo.org.nz;check puppet-nodes;OK;HARD;3;Active: active (running) since Mon 2022-06-06 07:57:42 UTC: 2min 49s ago
Jun  6 08:18:05 mgmt-b nagios3: SERVICE ALERT: db-b.foo.org.nz;check users-login;OK;HARD;3;USERS OK - 0 users currently logged in
Jun  6 08:19:55 mgmt-b nagios3: SERVICE ALERT: backup-b.foo.org.nz;check users-login;OK;HARD;3;USERS OK - 0 users currently logged in
Jun  6 08:20:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check users-login;OK;HARD;3;USERS OK - 0 users currently logged in
Jun  7 01:33:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check owncloud status;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 403 Forbidden - 454 bytes in 0.001 second response time
Jun  7 01:34:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check owncloud status;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 403 Forbidden - 454 bytes in 0.001 second response time
Jun  7 01:35:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check owncloud status;WARNING;HARD;3;HTTP WARNING: HTTP/1.1 403 Forbidden - 454 bytes in 0.001 second response time
Jun  7 01:45:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check users-login;WARNING;SOFT;1;USERS WARNING - 1 users currently logged in
Jun  7 01:46:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check users-login;WARNING;SOFT;2;USERS WARNING - 2 users currently logged in
Jun  7 01:47:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check users-login;WARNING;HARD;3;USERS WARNING - 2 users currently logged in
Jun  7 01:48:05 mgmt-b nagios3: SERVICE ALERT: db-b.foo.org.nz;check users-login;WARNING;SOFT;1;USERS WARNING - 1 users currently logged in
Jun  7 01:49:05 mgmt-b nagios3: SERVICE ALERT: db-b.foo.org.nz;check users-login;WARNING;SOFT;2;USERS WARNING - 1 users currently logged in
Jun  7 01:50:05 mgmt-b nagios3: SERVICE ALERT: db-b.foo.org.nz;check users-login;WARNING;HARD;3;USERS WARNING - 1 users currently logged in
Jun  7 02:10:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check owncloud status;OK;HARD;3;HTTP OK: HTTP/1.1 301 Moved Permanently - 535 bytes in 0.001 second response time
Jun  7 02:30:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check owncloud status;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 403 Forbidden - 454 bytes in 0.001 second response time
Jun  7 02:31:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check owncloud status;OK;SOFT;2;HTTP OK: HTTP/1.1 301 Moved Permanently - 707 bytes in 0.009 second response time
Jun  7 02:41:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check owncloud status;CRITICAL;SOFT;1;connect to address localhost and port 80: Connection refused
Jun  7 02:42:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check owncloud status;CRITICAL;SOFT;2;connect to address localhost and port 80: Connection refused
Jun  7 02:43:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check owncloud status;CRITICAL;HARD;3;connect to address localhost and port 80: Connection refused
Jun  7 02:48:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check owncloud status;OK;HARD;3;HTTP OK: HTTP/1.1 301 Moved Permanently - 707 bytes in 0.001 second response time
Jun  7 02:48:45 mgmt-b nagios3: SERVICE FLAPPING ALERT: app-b.foo.org.nz;check owncloud status;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
Jun  7 02:52:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check users-login;CRITICAL;HARD;3;USERS CRITICAL - 3 users currently logged in
Jun  7 03:13:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check owncloud status;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 404 Not Found - 451 bytes in 0.001 second response time
Jun  7 03:14:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check owncloud status;OK;SOFT;2;HTTP OK: HTTP/1.1 301 Moved Permanently - 707 bytes in 0.002 second response time
Jun  7 03:17:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check users-login;WARNING;HARD;3;USERS WARNING - 2 users currently logged in
Jun  7 03:24:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check owncloud status;CRITICAL;SOFT;1;connect to address localhost and port 80: Connection refused
Jun  7 03:25:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check owncloud status;CRITICAL;SOFT;2;connect to address localhost and port 80: Connection refused
Jun  7 03:26:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check owncloud status;CRITICAL;HARD;3;connect to address localhost and port 80: Connection refused
Jun  7 03:46:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check owncloud status;WARNING;HARD;3;HTTP WARNING: HTTP/1.1 404 Not Found - 451 bytes in 0.006 second response time
Jun  7 03:56:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check owncloud status;OK;HARD;3;HTTP OK: HTTP/1.1 301 Moved Permanently - 707 bytes in 0.001 second response time
Jun  7 05:26:45 mgmt-b nagios3: SERVICE FLAPPING ALERT: app-b.foo.org.nz;check owncloud status;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
Jun  7 05:42:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check users-login;OK;HARD;3;USERS OK - 0 users currently logged in
Jun  7 05:45:05 mgmt-b nagios3: SERVICE ALERT: db-b.foo.org.nz;check users-login;OK;HARD;3;USERS OK - 0 users currently logged in
Jun  7 09:27:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check users-login;WARNING;SOFT;1;USERS WARNING - 1 users currently logged in
Jun  7 09:28:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check users-login;WARNING;SOFT;2;USERS WARNING - 1 users currently logged in
Jun  7 09:29:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check users-login;WARNING;HARD;3;USERS WARNING - 1 users currently logged in
Jun  7 10:59:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check users-login;OK;HARD;3;USERS OK - 0 users currently logged in
Jun  7 11:29:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check users-login;WARNING;SOFT;1;USERS WARNING - 1 users currently logged in
Jun  7 11:30:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check users-login;WARNING;SOFT;2;USERS WARNING - 1 users currently logged in
Jun  7 11:31:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check users-login;WARNING;HARD;3;USERS WARNING - 1 users currently logged in
Jun  7 13:41:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check users-login;OK;HARD;3;USERS OK - 0 users currently logged in
Jun  8 00:21:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check users-login;WARNING;SOFT;1;USERS WARNING - 1 users currently logged in
Jun  8 00:22:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check users-login;WARNING;SOFT;2;USERS WARNING - 1 users currently logged in
Jun  8 00:23:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check users-login;WARNING;HARD;3;USERS WARNING - 1 users currently logged in
Jun  8 04:59:58 mgmt-b nagios3: SERVICE ALERT: backup-b.foo.org.nz;check users-login;WARNING;SOFT;1;USERS WARNING - 1 users currently logged in
Jun  8 05:00:08 mgmt-b nagios3: SERVICE ALERT: db-b.foo.org.nz;check users-login;WARNING;SOFT;1;USERS WARNING - 1 users currently logged in
Jun  8 05:00:58 mgmt-b nagios3: SERVICE ALERT: backup-b.foo.org.nz;check users-login;WARNING;SOFT;2;USERS WARNING - 1 users currently logged in
Jun  8 05:01:08 mgmt-b nagios3: SERVICE ALERT: db-b.foo.org.nz;check users-login;WARNING;SOFT;2;USERS WARNING - 1 users currently logged in
Jun  8 05:01:58 mgmt-b nagios3: SERVICE ALERT: backup-b.foo.org.nz;check users-login;WARNING;HARD;3;USERS WARNING - 1 users currently logged in
Jun  8 05:02:08 mgmt-b nagios3: SERVICE ALERT: db-b.foo.org.nz;check users-login;WARNING;HARD;3;USERS WARNING - 1 users currently logged in
Jun  8 09:16:58 mgmt-b nagios3: SERVICE ALERT: backup-b.foo.org.nz;check users-login;OK;HARD;3;USERS OK - 0 users currently logged in
Jun  8 09:17:08 mgmt-b nagios3: SERVICE ALERT: db-b.foo.org.nz;check users-login;OK;HARD;3;USERS OK - 0 users currently logged in
Jun  8 09:22:08 mgmt-b nagios3: SERVICE ALERT: db-b.foo.org.nz;check users-login;WARNING;SOFT;1;USERS WARNING - 1 users currently logged in
Jun  8 09:23:08 mgmt-b nagios3: SERVICE ALERT: db-b.foo.org.nz;check users-login;WARNING;SOFT;2;USERS WARNING - 1 users currently logged in
Jun  8 09:24:08 mgmt-b nagios3: SERVICE ALERT: db-b.foo.org.nz;check users-login;WARNING;HARD;3;USERS WARNING - 1 users currently logged in
Jun  8 10:58:48 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check users-login;OK;HARD;3;USERS OK - 0 users currently logged in
Jun  8 10:59:08 mgmt-b nagios3: SERVICE ALERT: db-b.foo.org.nz;check users-login;OK;HARD;3;USERS OK - 0 users currently logged in
Jun  9 03:53:48 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check users-login;WARNING;SOFT;1;USERS WARNING - 2 users currently logged in
Jun  9 03:54:48 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check users-login;WARNING;SOFT;2;USERS WARNING - 2 users currently logged in
Jun  9 03:55:48 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check users-login;WARNING;HARD;3;USERS WARNING - 2 users currently logged in
Jun  9 03:56:48 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check owncloud status;CRITICAL;SOFT;1;connect to address localhost and port 80: Connection refused
Jun  9 03:57:48 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check owncloud status;CRITICAL;SOFT;2;connect to address localhost and port 80: Connection refused
Jun  9 03:58:48 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check owncloud status;CRITICAL;HARD;3;connect to address localhost and port 80: Connection refused
Jun  9 04:25:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check users-login;CRITICAL;HARD;3;(null)
Jun  9 04:27:35 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;Remote disk servers;CRITICAL;SOFT;1;(null)
Jun  9 04:28:05 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check puppet-nodes;CRITICAL;SOFT;1;(null)
Jun  9 04:28:35 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;Remote disk servers;CRITICAL;SOFT;2;(null)
Jun  9 04:28:35 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check apache status;CRITICAL;SOFT;1;(null)
Jun  9 04:29:05 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check puppet-nodes;CRITICAL;SOFT;2;(null)
Jun  9 04:29:35 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;Remote disk servers;CRITICAL;HARD;3;(null)
Jun  9 04:29:36 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check apache status;CRITICAL;SOFT;2;(null)
Jun  9 04:29:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check hosts alive;CRITICAL;SOFT;1;(null)
Jun  9 04:30:05 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check puppet-nodes;CRITICAL;HARD;3;(null)
Jun  9 04:30:35 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check apache status;CRITICAL;HARD;3;(null)
Jun  9 04:30:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check hosts alive;CRITICAL;SOFT;2;(null)
Jun  9 04:31:45 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check hosts alive;CRITICAL;HARD;3;(null)
Jun  9 05:11:55 mgmt-b nagios3: SERVICE ALERT: backup-b.foo.org.nz;check users-login;WARNING;SOFT;1;USERS WARNING - 1 users currently logged in
Jun  9 05:12:55 mgmt-b nagios3: SERVICE ALERT: backup-b.foo.org.nz;check users-login;WARNING;SOFT;2;USERS WARNING - 1 users currently logged in
Jun  9 05:13:55 mgmt-b nagios3: SERVICE ALERT: backup-b.foo.org.nz;check users-login;WARNING;HARD;3;USERS WARNING - 1 users currently logged in
Jun  9 07:10:01 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check puppet-nodes;OK;HARD;3;Active: active (running) since Thu 2022-06-09 06:38:17 UTC: 31min ago
Jun  9 07:10:41 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check apache status;OK;HARD;3;Active: active (running) since Thu 2022-06-09 05:51:39 UTC: 1h 18min ago
Jun  9 07:10:51 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check users-login;WARNING;HARD;3;USERS WARNING - 2 users currently logged in
Jun  9 07:11:41 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check hosts alive;OK;HARD;3;PING OK - Packet loss = 0%, RTA = 0.05 ms
Jun  9 07:13:51 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check owncloud status;OK;HARD;3;HTTP OK: HTTP/1.1 301 Moved Permanently - 707 bytes in 0.001 second response time
Jun  9 07:14:41 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;Remote disk servers;OK;HARD;3;DISK OK - free space: /var/tmp 23886 MB (80% inode=94%):
Jun  9 08:14:01 mgmt-b nagios3: SERVICE ALERT: backup-b.foo.org.nz;check users-login;OK;HARD;3;USERS OK - 0 users currently logged in
Jun  9 08:15:51 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check users-login;OK;HARD;3;USERS OK - 0 users currently logged in
Jun 10 01:58:51 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check owncloud status;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 403 Forbidden - 454 bytes in 0.001 second response time
Jun 10 01:59:51 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check owncloud status;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 403 Forbidden - 454 bytes in 0.001 second response time
Jun 10 02:00:51 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check owncloud status;WARNING;HARD;3;HTTP WARNING: HTTP/1.1 403 Forbidden - 454 bytes in 0.001 second response time
Jun 10 02:10:51 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check users-login;WARNING;SOFT;1;USERS WARNING - 1 users currently logged in
Jun 10 02:11:51 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check users-login;WARNING;SOFT;2;USERS WARNING - 1 users currently logged in
Jun 10 02:12:51 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check users-login;WARNING;HARD;3;USERS WARNING - 1 users currently logged in
Jun 10 02:19:11 mgmt-b nagios3: SERVICE ALERT: db-b.foo.org.nz;check users-login;WARNING;SOFT;1;USERS WARNING - 1 users currently logged in
Jun 10 02:20:11 mgmt-b nagios3: SERVICE ALERT: db-b.foo.org.nz;check users-login;WARNING;SOFT;2;USERS WARNING - 1 users currently logged in
Jun 10 02:21:11 mgmt-b nagios3: SERVICE ALERT: db-b.foo.org.nz;check users-login;WARNING;HARD;3;USERS WARNING - 1 users currently logged in
Jun 10 02:25:51 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check owncloud status;CRITICAL;HARD;3;connect to address localhost and port 80: Connection refused
Jun 10 02:30:51 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check owncloud status;WARNING;HARD;3;HTTP WARNING: HTTP/1.1 403 Forbidden - 454 bytes in 0.001 second response time
Jun 10 02:35:51 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check owncloud status;OK;HARD;3;HTTP OK: HTTP/1.1 301 Moved Permanently - 535 bytes in 0.001 second response time
Jun 10 02:35:51 mgmt-b nagios3: SERVICE FLAPPING ALERT: app-b.foo.org.nz;check owncloud status;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
Jun 10 03:27:51 mgmt-b nagios3: SERVICE ALERT: app-b.foo.org.nz;check users-login;OK;HARD;3;USERS OK - 0 users currently logged in
Jun 10 03:31:11 mgmt-b nagios3: SERVICE ALERT: db-b.foo.org.nz;check users-login;OK;HARD;3;USERS OK - 0 users currently logged in
Jun 10 04:10:51 mgmt-b nagios3: SERVICE FLAPPING ALERT: app-b.foo.org.nz;check owncloud status;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)